# Tarea 01 — Funciones básicas

Esta tarea tiene 2 ejercicios. Lee con atención cada enunciado y completa las celdas marcadas con `# YOUR CODE HERE`.

**Importante:** no modifiques las celdas de tests (las que empiezan con `# Hidden tests` se calificarán automáticamente al entregar).

Antes de entregar, usa el botón **Validate** para verificar que tu código pasa los tests visibles.

## Ejercicio 1 (3 puntos)

Implementa una función `factorial(n)` que dado un entero no negativo `n`, devuelva `n!`.

Recuerda que `0! = 1` y `n! = n * (n-1)!`.


<!-- grade_id: ej1-solucion, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    ### BEGIN SOLUTION
    if n < 0:
        raise ValueError("n debe ser no negativo")
    result = 1
    for i in range(2, n + 1):
        result *= i
    return result
    ### END SOLUTION
```

<!-- grade_id: ej1-tests-visibles, points: 1, grade: true, solution: false -->
```python
# Tests visibles (los alumnos los ven)
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(5) == 120
```

<!-- grade_id: ej1-tests-ocultos, points: 2, grade: true, solution: false -->
```python
# Hidden tests — los alumnos NO los ven en su versión
### BEGIN HIDDEN TESTS
assert factorial(10) == 3628800
assert factorial(7) == 5040
try:
    factorial(-1)
    assert False, 'Debió lanzar ValueError'
except ValueError:
    pass
### END HIDDEN TESTS
```

## Ejercicio 2 (2 puntos)

Implementa `es_primo(n)` que devuelve `True` si `n` es un número primo, `False` en caso contrario.

Considera que los números menores que 2 NO son primos.

<!-- grade_id: ej2-solucion, points: 0, grade: false, solution: true -->
```python
def es_primo(n):
    ### BEGIN SOLUTION
    if n < 2:
        return False
    for i in range(2, int(n**0.5) + 1):
        if n % i == 0:
            return False
    return True
    ### END SOLUTION
```

<!-- grade_id: ej2-tests, points: 2, grade: true, solution: false -->
```python
assert es_primo(2) == True
assert es_primo(7) == True
assert es_primo(1) == False
assert es_primo(15) == False
### BEGIN HIDDEN TESTS
assert es_primo(0) == False
assert es_primo(-5) == False
assert es_primo(97) == True
assert es_primo(100) == False
### END HIDDEN TESTS
```